# DO NOT ADD OR CREATE ANY NEW CELLS
## DO NOT ADD OR CREATE ANY NEW CELLS
### DO NOT ADD OR CREATE ANY NEW CELLS
#### DO NOT ADD OR CREATE ANY NEW CELLS

### YOU HAVE TO RUN ALL THESE CELLS BEFORE MOVING ON TO YOUR QUESTIONS. IF YOU ENCOUNTER ANY ERRORS JUST RESTART THE KERNEL

```python
#Do not modify this cell
#Run this Cell
import pandas as pd
%load_ext sql
```

```python
q1 = pd.read_csv('q1.csv', sep = '\t')
q2 = pd.read_csv('q2.csv', sep = '\t')
q3 = pd.read_csv('q3.csv', sep = '\t')
q4 = pd.read_csv('q4.csv', sep = '\t')
q5 = pd.read_csv('q5.csv', sep = '\t')
```

Output:
```
File ~\miniconda3\Lib\site-packages\pandas\io\parsers\readers.py:1880, in TextFileReader._make_engine(self, f, engine)
   1878     if "b" not in mode:
   1879         mode += "b"
-> 1880 self.handles = get_handle(
   1881     f,
   1882     mode,
   1883     encoding=self.options.get("encoding", None),
   1884     compression=self.options.get("compression", None),
   1885     memory_map=self.options.get("memory_map", False),
   1886     is_text=is_text,
   1887     errors=self.options.get("encoding_errors", "strict"),
   1888     storage_options=self.options.get("storage_options", None),
   1889 )
   1890 assert self.handles is not None
   1891 f = self.handles.handle

File ~\miniconda3\Lib\site-packages\pandas\io\common.py:873, in get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text, errors, storage_options)
    868 elif isinstance(handle, str):
    869     # Check whether the filename is to be opened in binary mode.
    870     # Binary mode does not support 'enco
```

```python
%sql sqlite:///bookstore.sqlite
```

# Here are your questions
You are the database administrator for an online bookstore. The bookstore database has several tables that store information about books, authors, customers, orders, and order items. Your task is to create a set of SQL queries to extract useful information from this database.
- The name of the table are authors, books, customers, order_items, orders, reviews
- Always start your answers with this line `%%sql result <<`. Each question has its own format. For example, Question One would be `%%sql q1_result <<` and so on like that.
- Then start writing your query on the next line for example, Question 2 would be
    ```
%%sql q2_result <<
    select * from orders
    ```

- If you do not follow these instructions your score will not count.

#### 1. Retrieve a list of the top 5 most popular books based on the number of times they have been ordered. Also order by author name in descending order

```python
#Run this cell to see the expected output
#do not modify this cell
print('This is the expected output')
q1
```

Output:
```
This is the expected output
NameError                                 Traceback (most recent call last)
Cell In[3], line 4
      1 #Run this cell to see the expected output
      2 #do not modify this cell
      3 print('This is the expected output')
----> 4 q1

NameError: name 'q1' is not defined
```

- **Your column names must match the expected output**
- **clear everything in the cell below and write your sql query there.** 
- Start your answer by typing this exactly as the first line.
```
%%sql q1_result <<


<!-- grade_id: cell-e0bd2ff299f43c54, points: 0, grade: false, solution: true -->
```python
%%sql q1_result <<
Select b.title, a.name as author_name, Count(oi.book_id) as order_count
from books as b
Join authors as a
on b.author_id = a.author_id
Join order_items as oi
On b.book_id = oi.book_id
GROUP BY b.title
ORDER BY order_count DESC, author_name DESC
LIMIT 5;
```

Output:
```
Traceback (most recent call last):
  File "C:\Users\USER\miniconda3\Lib\site-packages\sql\magic.py", line 196, in execute
    conn = sql.connection.Connection.set(
           ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
  File "C:\Users\USER\miniconda3\Lib\site-packages\sql\connection.py", line 82, in set
    raise ConnectionError(
sql.connection.ConnectionError: Environment variable $DATABASE_URL not set, and no connect string given.

Connection info needed in SQLAlchemy format, example:
               postgresql://username:password@hostname/dbname
               or an existing connection: dict_keys([])
```

```python
#Run this cell to check your output
#do not modify this cell
print('This is the output of your query')
q1_result.DataFrame()
```

Output:
```
This is the output of your query
NameError                                 Traceback (most recent call last)
Cell In[5], line 4
      1 #Run this cell to check your output
      2 #do not modify this cell
      3 print('This is the output of your query')
----> 4 q1_result.DataFrame()

NameError: name 'q1_result' is not defined
```

<!-- grade_id: cell-cf9ac61b23e1045c, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
assert q1_result.DataFrame().equals(q1)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[6], line 2
      1 #do not modify this cell
----> 2 assert q1_result.DataFrame().equals(q1)

NameError: name 'q1_result' is not defined
```

#### 2. Calculate the total revenue generated by each author from the sales of their books.

```python
#Run this cell to see the expected output
#do not modify this cell
print('This is the expected output')
q2['total_revenue'] = q2['total_revenue'].apply(lambda x: round(x, 2) )
q2
```

Output:
```
This is the expected output
NameError                                 Traceback (most recent call last)
Cell In[6], line 4
      1 #Run this cell to see the expected output
      2 #do not modify this cell
      3 print('This is the expected output')
----> 4 q2['total_revenue'] = q2['total_revenue'].apply(lambda x: round(x, 2) )
      5 q2

NameError: name 'q2' is not defined
```

- **Your column names must match the expected output**
- **clear everything in the cell below and write your sql query there.** 
- Start your answer by typing this exactly as the first line.
```
%%sql q2_result <<


<!-- grade_id: cell-53a525f69f919b9e, points: 0, grade: false, solution: true -->
```python
%%sql q2_result <<
Select a.name as author_name,
sum(b.price * oi.quantity) as total_revenue
FROM order_items oi
Join books b on oi.book_id = b.book_id
Join authors a ON b.author_id =a.author_id
GROUP BY author_name, a.author_id
Order by total_revenue DESC;
```

Output:
```
Traceback (most recent call last):
  File "C:\Users\USER\miniconda3\Lib\site-packages\sql\magic.py", line 196, in execute
    conn = sql.connection.Connection.set(
           ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
  File "C:\Users\USER\miniconda3\Lib\site-packages\sql\connection.py", line 82, in set
    raise ConnectionError(
sql.connection.ConnectionError: Environment variable $DATABASE_URL not set, and no connect string given.

Connection info needed in SQLAlchemy format, example:
               postgresql://username:password@hostname/dbname
               or an existing connection: dict_keys([])
```

```python
#Run this cell to check your output
#do not modify this cell
print('This is the output of your query')
q2result = q2_result.DataFrame()
q2result['total_revenue'] = q2result['total_revenue'].apply(lambda x: round(x, 2) )
q2result
```

Output:
```
This is the output of your query
NameError                                 Traceback (most recent call last)
Cell In[8], line 4
      1 #Run this cell to check your output
      2 #do not modify this cell
      3 print('This is the output of your query')
----> 4 q2result = q2_result.DataFrame()
      5 q2result['total_revenue'] = q2result['total_revenue'].apply(lambda x: round(x, 2) )
      6 q2result

NameError: name 'q2_result' is not defined
```

<!-- grade_id: cell-659def203a1c50ec, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
assert q2result.equals(q2)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[9], line 2
      1 #do not modify this cell
----> 2 assert q2result.equals(q2)

NameError: name 'q2result' is not defined
```

#### 3. Identify the top 5 customers who have placed the most orders. Also order by customer name in descending order

```python
#Run this cell to see the expected output
#do not modify this cell
print('This is the expected output')
q3
```

Output:
```
This is the expected output
NameError                                 Traceback (most recent call last)
Cell In[11], line 4
      1 #Run this cell to see the expected output
      2 #do not modify this cell
      3 print('This is the expected output')
----> 4 q3

NameError: name 'q3' is not defined
```

- **Your column names must match the expected output**
- **clear everything in the cell below and write your sql query there.** 
- Start your answer by typing this exactly as the first line.
```
%%sql q3_result <<

<!-- grade_id: cell-152c04c839ff086b, points: 0, grade: false, solution: true -->
```python
%%sql q3_result <<
SELECT c.name as customer_name, count(o.order_id) as order_count
FROM customers c
JOIN orders o ON  c.customer_id = o.customer_id
GROUP BY c.name
ORDER BY order_count DESC, customer_name DESC
LIMIT 5;

```

Output:
```
Traceback (most recent call last):
  File "C:\Users\USER\miniconda3\Lib\site-packages\sql\magic.py", line 196, in execute
    conn = sql.connection.Connection.set(
           ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
  File "C:\Users\USER\miniconda3\Lib\site-packages\sql\connection.py", line 82, in set
    raise ConnectionError(
sql.connection.ConnectionError: Environment variable $DATABASE_URL not set, and no connect string given.

Connection info needed in SQLAlchemy format, example:
               postgresql://username:password@hostname/dbname
               or an existing connection: dict_keys([])
```

```python
#Run this cell to check your output
#do not modify this cell
print('This is the output of your query')
q3_result.DataFrame()
```

Output:
```
This is the output of your query
NameError                                 Traceback (most recent call last)
Cell In[13], line 4
      1 #Run this cell to check your output
      2 #do not modify this cell
      3 print('This is the output of your query')
----> 4 q3_result.DataFrame()

NameError: name 'q3_result' is not defined
```

<!-- grade_id: cell-9aedb87130b9508f, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
assert q3_result.DataFrame().equals(q3)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[14], line 2
      1 #do not modify this cell
----> 2 assert q3_result.DataFrame().equals(q3)

NameError: name 'q3_result' is not defined
```

#### 4. Find all orders that contain book with book_id = 1.

```python
#Run this cell to check your output
#do not modify this cell
print('This is the expected output')
q4
```

Output:
```
This is the expected output
NameError                                 Traceback (most recent call last)
Cell In[15], line 4
      1 #Run this cell to check your output
      2 #do not modify this cell
      3 print('This is the expected output')
----> 4 q4

NameError: name 'q4' is not defined
```

- **Your column names must match the expected output**
- **clear everything in the cell below and write your sql query there.** 
- Start your answer by typing this exactly as the first line.
```
%%sql q4_result <<

<!-- grade_id: cell-43cea19331ed2979, points: 0, grade: false, solution: true -->
```python
%%sql q4_result <<
Select oi.order_id,
c.name as customer_name,
o.order_date from order_items oi
JOIN orders o ON oi.order_id = o.order_id
JOIN customers c ON o.customer_id = c.customer_id
Where oi.book_id = 1
```

Output:
```
* sqlite:///bookstore.sqlite
Done.
Returning data to local variable q4_result
```

```python
#Run this cell to check your output
#do not modify this cell
print('This is the output of your query')
q4_result.DataFrame()
```

Output:
```
This is the output of your query
   order_id  customer_name  order_date
0         1  Alice Johnson  2023-05-01
1         5   Eva Martinez  2023-05-05
2        13   Eva Martinez  2023-05-15
3        21  Charlie Brown  2023-06-20
```

<!-- grade_id: cell-0103803ea2a61006, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
assert q4_result.DataFrame().equals(q4)
```

#### 5. Select detailed review information of the Least 5 Books based on ratings, including book title, author name, review rating, and review text.

```python
#Run this cell to see the expected output
#do not modify this cell
print('This is the expected output')
q5
```

Output:
```
This is the expected output
   review_id                               book_title     author_name  rating  \
0          1                                     1984   George Orwell       5   
1          2                      Pride and Prejudice     Jane Austen       4   
2          3  Harry Potter and the Philosophers Stone    J.K. Rowling       5   
3          4           Adventures of Huckleberry Finn      Mark Twain       4   
4          5                               The Hobbit  J.R.R. Tolkien       5   

                                      review_text  
0                         A timeless masterpiece.  
1                         A wonderful love story.  
2  An incredible journey of magic and friendship.  
3                        A fascinating adventure.  
4                        An epic tale of heroism.
```

- **Your column names must match the expected output**
- **clear everything in the cell below and write your sql query there.** 
- Start your answer by typing this exactly.
```
%%sql q5_result <<

<!-- grade_id: cell-e51c13847d1d0ebf, points: 0, grade: false, solution: true -->
```python
%%sql q5_result <<
Select r.review_id,
b.title as book_title,
a.name as author_name,
r.rating,
r.review_text
FROM reviews r
JOIN books b ON r.book_id =b.book_id
JOIN authors a ON b.author_id =a.author_id
ORDER BY r.review_id ASC
LIMIT 5;
```

Output:
```
* sqlite:///bookstore.sqlite
Done.
Returning data to local variable q5_result
```

```python
#Run this cell to check your output
#do not modify this cell
print('This is the output of your query')
q5_result.DataFrame()
```

Output:
```
This is the output of your query
   review_id                               book_title     author_name  rating  \
0          1                                     1984   George Orwell       5   
1          2                      Pride and Prejudice     Jane Austen       4   
2          3  Harry Potter and the Philosophers Stone    J.K. Rowling       5   
3          4           Adventures of Huckleberry Finn      Mark Twain       4   
4          5                               The Hobbit  J.R.R. Tolkien       5   

                                      review_text  
0                         A timeless masterpiece.  
1                         A wonderful love story.  
2  An incredible journey of magic and friendship.  
3                        A fascinating adventure.  
4                        An epic tale of heroism.
```

<!-- grade_id: cell-bb340718cdf0c063, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
assert q5_result.DataFrame().equals(q5)
```